# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the play button to run the code cell below.

```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment05.ipynb")
```

# Assignment 05: User-Defined Functions, Control and Iteration

Welcome to Assignment 05!  Throughout the course you will complete assignments like this one. You can't learn technical subjects without hands-on practice, so these assignments are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on a question, so ask a post to the discussion board or ask your instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

To receive credit for this assignment, answer all questions correctly and submit before the deadline.

**Due Date:** Monday, July 11, 2022 @ 11:59 pm

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## Today's Assignment

In today's assignment, you'll learn how to:

- read error messages.

- shutdown and restart the kernel of a Jupyter notebook.

- define a function.

- use `if` statements.

- use `for` loops.

Let's get started! Run the cell below.

```python
from datascience import *
import numpy as np
import math
```

## Errors

Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:

1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.

2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make many errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

<!-- END QUESTION -->



**Note:** In the toolbar, there is the option to click `Run > Run All Cells`, which will run all the code cells in this notebook in order. However, the notebook stops running code cells if it hits an error, like the one in the cell above.

You should see something like this (minus our annotations):

<img src="images/error.jpg" />

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages, but you don't need to know it all in order to program effectively. If you see a cryptic message like this, you can often get by without deciphering it. Of course, if you're frustrated, ask a neighbor, or your instructor for help.

Try to fix the code above so that you can run the cell and see the intended message instead of an error.

## The Kernel

The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

Next to every code cell, you'll see some text that says `In [...]`. Before you run the cell, you'll see `In [ ]`. When the cell is running, you'll see `In [*]`. If you see an asterisk (\*) next to a cell that doesn't go away, it's likely that the code inside the cell is taking too long to run, and it might be a good time to interrupt the kernel (discussed below). When a cell is finished running, you'll see a number inside the brackets, like so: `In [1]`. The number corresponds to the order in which you run the cells; so, the first cell you run will show a 1 when it's finished running, the second will show a 2, and so on. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:

1. At the top of your screen, click **Kernel**, then **Interrupt Kernel**. Trying running your code again.

2. If that doesn't help, click **Kernel**, then **Restart Kernel**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.

3. If that doesn't help, restart your server. This is very rarely needed, but it's good to know how to do just in case. First, save your work by clicking **File** at the top left of your screen, then **Save Notebook**. Next, click **File** again and choose **Hub Control Panel** and select **Stop My Server** to shut it down. Then, after a few moments select **Start My Server** to start it back up. Then, navigate back to the notebook you were working on. You'll still have to run your code cells again.

## Functions

Functions allow us to combine the topics we have discussed so far in the course. Writing functions is useful when we want to use the same code repeatedly without wanting to write it out each and every time. You can think of functions like variables, they store lines of code instead of Python values, and they can be used when you need them.

Below is an example of a Python funciton that doubles its input:

```python
def double(x):
    """Doubles the value of the input."""
    return 2 * x
```

Let's break down the different parts of this function:

- The `def` at the beginning tells Python that we are making a new function.

- The text in between `"""` and `"""` is called a [`docstring`](https://www.geeksforgeeks.org/python-docstrings/). Docstrings provide a convenient way of documenting Python functions.

- The blue `double` text is the name of our function. If we want to use this function, we will call it using this name.

- The `x` inside the parentheses is the input variable, otherwise known as a function's domain. A function can have as many input variables as you need, as long as they have different names.

- The `:` indicates we are done defining our function's name and input variables.

- All lines after this `def` line are **indented** to indicate we are is what we call the **body** of the function.

    - You can have as many lines as you want in your functions, but in this case we only have 1.


- The `return` keyword indicates we want to output the following evaluated Python expression. In our case, we want to multiply our input variable by 2.

Let's see what happens when we use our function:

```python
answer = double(5)
# We expect the value of this 'answer' variable to be 10 as long as our function is working correctly
answer
```

**Question 1.** Write a function that converts from **feet** to **meters**.

**Hint:** Use the conversion 1 foot = 0.3048 meters.

```python
def feet_to_meters(x):
    """Converts input from feet to meters."""
    return ...
feet_to_meters(10)
```

```python
grader.check("q1")
```

## Should I Stay or Should I Go

Now let's take a look at a function that uses **variables** and **boolean operators**. Read through the functions below and see if you understand the code. In particular, make sure you understand the assignment statements

```
green = (light == 'green')
yellow = (light == 'yellow')
red = (light == 'red')
```

and the comparison statement

```
return (green or yellow) and (not red)
```

After you understand the code explain it to a friend, then explain it to your instructor. 

```python
# Based on the color of the stoplight in front of us we want to know whether we should stop or keep going.
# Because we know we are going to encounter many lights on our trip, we should write a function.

# We can make 2 functions: one where we take risks and go on yellow lights,
# and one where we play it safe and stop for yellow lights

def stop_or_go_risky(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return (green or yellow) and (not red) # Go on green and yellow, stop on red

def stop_or_go_safe(light):
    green = (light == 'green')
    yellow = (light == 'yellow')
    red = (light == 'red')
    return green and ((not yellow) and (not red)) # Go on green, stop on yellow and red

# These functions are booleans because they output either True or False. We can write functions 
# that output any type we need, and you will do so throughout the course.
```

Run the following cell as a demo:

```python
# We can use the outputs of these functions to let us know to stop the car 
# if we need to on our way to the market.

light1 = stop_or_go_safe('green')
light2 = stop_or_go_safe('yellow')
light3 = stop_or_go_safe('red')

print("I'm driving to the market. No rush.")
print("") # This is the same as hitting the return key while typing in Google Docs
print("Should I go at a green light?... " + str(light1))
print("Should I go at a yellow light?... " + str(light2))
print("Should I go at a red light?... " + str(light3))
print("")

# But if we find out we are running late, we may want to take more risks.
print("Oh no! I'm running late!")
print("")

light4 = stop_or_go_risky('green')
light5 = stop_or_go_risky('yellow')
light6 = stop_or_go_risky('red')

print("Should I go at a green light?... " + str(light4))
print("Should I go at a yellow light?... " + str(light5))
print("Should I go at a red light?... " + str(light6))
```

**Question 2.** After you explain the code to your instructor assign the value of `"yes"` to the variable `checked_by_the_instructor`

```python
checked_by_the_instructor = ...
checked_by_the_instructor
```

```python
grader.check("q2")
```

# Control

When writing functions in Python, we may want the function to behave differently depending on the input. We could choose to write several similar functions to accomplish this, but that would require copying and pasting much of the same code over and over again, only making small changes.

Instead, we need a way to tell one singular function to execute different code for different inputs. In lecture we discussed this idea as `if-else` statements.

Let's look at an example of a function that uses `if-else` statements to tell us what water looks like at different temperatures:

<img src='images/water.jpeg' width=300>

```python
def state_of_water(temperature):
    """Print the state of water based on the input 
       of the temperature."""
    if temperature <= 32: # Water is solid at and below 32°F
        return "ice"
    elif temperature < 212: # Water is liquid between 32°F and 212°F
        return "liquid"
    else: # Water is gaseous above 212°F
        return "steam"
```

Here, we could have made 3 different functions that each dealt with a different state of water, but that would force us to know the state of water (what we currently do not know and are trying to figure out) so that we could pick the right function to use. This is why control such as `if-else` statements are so important!

Now let's use this function to determine the state of water in the following locations during the winter:

| **City**     | Temperature (°F) |
| -------------| ---------------- |
| Durham       | 58 |
| New York     | 24 |
| Miami        | 78 |
| Earth's Core | 10800 |

We can see the calculation each city's state of water using the function defined above:

```python
print("Water in Durham is: " + state_of_water(58))
```

```python
print("Water in New York is: " + state_of_water(24))
```

```python
print("Water in Miami is: " + state_of_water(78))
```

```python
print("Water at the Center of the Earth is: " + state_of_water(10800))
```

We can call the function on different inputs and the function deals with each input differently based on the control logic you gave it. This will be very important when you write functions that deal with large, sometimes unpredictable datasets where the logic of your function will deal with inputs you may not have directly prepared the function for.

**Question 3.** Write a function that prints the letter grade based on the table below.

| **Letter Grade** | **Average** |
|:-----------------|:----------- |
| A | 90 - 100     |
| B | 80 - 89      |
| C | 70 - 79      |
| D | 60 - 69      |
| F | less than 60 |

```python
def letter_grade(average):
    """Prints the letter grade based on the input
       of the average as a whole number."""
    if ... :
       return ...
    elif ... :
       return ...
    ...
letter_grade(89)
```

```python
grader.check("q3")
```

## Iteration

Iteration is the repetition of code. Typically, the variables in iteration change sequentially, and the iteration continues to repeat until a boolean condition indicates that the iteration be stopped. 

Iteration is a very important tool in Python. `For` loops are a form of iteration that loop over a Python data structure, such as a array or a string. 

As an example, we can iterate over numbers using the `range()` function. Run the cell below to see an example.

```python
for i in range(3):
    print(i)
```

Notice that the first value is 0 and the last value is 2. If we wanted to print

```
1

2

3
```

we could do this.

```python
for i in range(1,4,1):
    print(i)
```

Let's breakdown the numbers inside the parenthesis of the `range()` function.

`
range(<start_value>, <end_value>, <increment_value>)
`

- 1 - The first value is the start value.

- 4 - The second value is the end value (notice that the `range()` function ends prior to the end value).
- 1 - The third value is the increment value (if no value is specified, then the increment will default to 1).

Run the cell below to see what happens when no increment value is specified.

```python
for i in range(1,4):
    print(i)
```

Another fun example.

```python
for i in range(10):
    print("Charging... battery at:", str(i * 10), '%')
print("Battery charged at 100%!")
```

Notice how the value of `i` changes on each loop because of the `range()` function.

**Important:** We do not necessarily have to use `i` as the name of our `for` loop variable. We can use any name we want.

We can iterate over the characters in a string!

```python
string = "Data Science is Cool!"
for letter in string: # 'letter' instead of 'i'
    print(letter)
```

We can even iterate over the words in a string.

```python
sentence = string.split(" ") # Split up string, using spaces (" ") to separate words
for word in sentence: # 'word' instead of 'i'
    print(word)
```

We can also iterate over an array.

```python
even_digits = make_array(2, 4, 6, 8)
even_digits
```

```python
for digit in even_digits: # 'digit' instead of 'i'
    print(digit)
```

<!-- BEGIN QUESTION -->

**Question 4.** Write a `for` loop to print the cube of all numbers from 1 up to and including a given number.

**Hint:** If your code is correct the output should look like this (that is, if `given_number` is assigned to 5, `given_number=5`).

```
The current number is : 1  and the cube is 1
The current number is : 2  and the cube is 8
The current number is : 3  and the cube is 27
The current number is : 4  and the cube is 64
The current number is : 5  and the cube is 125
```

```python
given_number = 5
for ... in ...:
    print("The current number is :", i, " and the cube is", ...)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.** Write a `for` loop to print all the even numbers from 1 up to and including a given number.

**Hint:** If your code is correct the output should look like this (that is, if `given_number` is assigned to 10, `given_number=10`).

```
The number  2  is even
The number  4  is even
The number  6  is even
The number  8  is even
The number  10  is even
```

```python
given_number = 10
for ... in ...:
    if ...
        print("The number ", i, " is even")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**Question 5.** Write a `for` loop to print all the numbers from 1 up to and including a given number. If the number is even print 

```
The number  2  is even
```

but if the number is odd print

```
The number  1  is odd
```

**Hint:** If your code is correct the output should look like this (that is, if `given_number` is assigned to 10, `given_number=10`).

```
The number  1  is odd
The number  2  is even
The number  3  is odd
The number  4  is even
The number  5  is odd
The number  6  is even
The number  7  is odd
The number  8  is even
The number  9  is odd
The number  10  is even
```

```python
given_number = 10
for ... in ...:
    if ...
        print("The number ", i, " is even")
    ...
        print("The number ", i, " is odd")
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by `SHIFT`-clicking on the file name and selecting **Save Link As**. Or, find the .zip file in the left side of the screen and right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: None)
```
>>> feet_to_meters(10) == 3.048
True
```


Tests for q2 (points: None)
```
>>> checked_by_the_instructor
'yes'
```


Tests for q3 (points: None)
```
>>> letter_grade(90) == 'A'
True
```
```
>>> letter_grade(83) == 'B'
True
```
```
>>> letter_grade(73) == 'C'
True
```
```
>>> letter_grade(63) == 'D'
True
```
```
>>> letter_grade(53) == 'F'
True
```
