# Run the cell below

To run a code cell (i.e.; execute the python code inside a Jupyter notebook) you can click the play button on the ribbon underneath the name of the notebook. Before you begin click the play button to run the code cell below.

```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment02.ipynb")
```

# Assignment 02: Variables and Comparison Operators

Welcome to Assignment 02!  Throughout the course you will complete assignments like this one. You can't learn technical subjects without hands-on practice, so these assignments are an important part of the course.

Collaborating on labs is more than okay -- it's encouraged! You should rarely remain stuck for more than a few minutes on a question, so ask a post to the discussion board or ask your instructor for help. Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it. You should **not** just copy/paste someone else's code, but rather work together to gain understanding of the task you need to complete. 

To receive credit for this assignment, answer all questions correctly and submit before the deadline.

**Due Date:** Thursday, July 7, 2022 @ 11:59 pm

**Collaboration Policy:** Data science is a collaborative activity. While you may talk with others about the labs, we ask that you **write your solutions individually**. If you do discuss the assignments with others **please include their names below** (it's a good way to learn your classmates' names).

**Collaborators:** 

List collaborators here.

## Today's Assignment

In today's assignment, you'll learn how to:

- determine variable types.

- cast variable types.

- use comparison operators.

Let's get started!

## Variable Types

Variables in Python can have many different types, but we are most interested in focusing on the following:

- Integers (`int`)
    - Integers are whole numbers, positive or negative, without decimals.


- Strings/Text (`str`)
    - Strings are sequences of Unicode characters.
    - **Note:** Click [here](https://youtu.be/EGtcgMlyBhU) to see a video that explains what Unicode is, why it was developed, and how computers use it to display the characters we see on our screens.


- Floats (`float`)    
    - Floats are numbers that are not integers (decimals and fractions mostly).


- Booleans (`bool`)
    - Booleans can only be `True` or `False`.


- NoneType (`None`)
    - Only the value `None`; indicates null value for a variable

<!-- BEGIN QUESTION -->

**Question 1.** Briefly (in 3-5 sentences) describe what Unicode is wand why it was developed.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

Run the following code cells to see what these variable types look like in Python:

```python
i_am_an_integer = 7
type(i_am_an_integer)
```

```python
i_am_a_string = "hello"
type(i_am_a_string)
```

```python
i_am_a_float = 4.0
type(i_am_a_float)
```

```python
i_am_a_boolean = True
type(i_am_a_boolean)
```

```python
i_am_a_nonetype = None
type(i_am_a_nonetype)
```

As we have seen, we can use these values in other Python expressions. Although, we have to be careful when doing Python operations on variables of different types.

Different types are often not compatible with each other as in the following code cell:

```python
# DO NOT CHANGE
i_am_an_integer + i_am_a_string
```

You can fix the error above by **casting** the variable `i_am_an_integer` to a string. The `str()` function will change a numerical variable (`int` or `float`) to a string. The `+` operator concatenates, or joins two strings together.

Run the cell below to see an example.

```python
# DO NOT CHANGE
str(i_am_an_integer) + i_am_a_string
```

## Variable Type Casting

In data science, it is common to have data in a table that represents numbers, but the table stores these numbers as strings (`"1"` instead of `1`). In this case, Python has built-in functions intended to change the type of the variable as long as the variable can validly be represented as the type you are changing to. This is called **casting** from one type to another.

You can change `"1"` to the integer `1`, but you cannot change `"hello"` into a valid integer.

```python
int("1")
```

```python
int("hello")
```

Read the paragraph below to learn about the context of the question.

Regulations of vapes and electronic cigarettes are complex and vary across countries. Some countries have no regulations, while others ban vapes completely. As of Dec. 20, 2019, federal law prohibits retailers from selling tobacco products (cigarettes, cigars, and e-cigarettes) and vape products to people under 21 years old. As State law enforcement, NC Alcohol Law Enforcement (ALE) will continue to enforce North Carolina law which prohibits the sale of tobacco products to anyone under 18 years old. Any sale of tobacco products to someone under 21 may be subject to federal authorities’ enforcement. 

In early September 2019, Michigan became the first state to announce its intent to limit the sale of vaping products, when Gov. Gretchen Whitmer said she would issue an emergency ban on the online and retail sale of nicotine vaping products in any flavor except tobacco. Whitmer also said she would restrict the marketing of vaping products by forbidding the use of terms like “clean,” “safe” and “healthy.”

In Canada, it is illegal to sell or provide vaping products to anyone under the age of 18. Although some provincial laws have increased this age to 19 or 21.

**Question 2.** The **casting** relationship goes both ways. Say your table has numbers stored as integers, but you want to print out a sentence that contains those numbers. Cast the `nc_age` and `ca_age` variables to another type so that the following cell does not error, and in fact prints out the sentence you wish to see. 

```python
nc_age = 21 # DO NOT CHANGE
ca_age = 18 # DO NOT CHANGE

print("The minimum age to buy e-cigarette in the US is " + us_age
     + ", but in Canada the minimum age is " + ca_age + ".")
```

## Boolean Operators

### Comparison operators

Comparison operators are boolean operators that are used to compare two different values in Python.

We introduce 6 of these operators in lecture: `==`, `!=`,`<`, `<=`, `>`, and `>=`. Try using these operators on variables and values of all types, though in practice they are typically used with integers and floats, and sometimes on strings. Strings are treated in increasing alphabetical order: ('a' < 'b') == True.

Run the following series of cells to see some examples of these comparison operators in practice:

```python
1 < 2
```

```python
min(20, 30) <= max(-10, 20)
```

```python
(3 + 4) != (3 * 4)
```

```python
"abc" < "def"
```

There is an important distinction here between `=` and `==`:

- `=` is used to **assign** values to variables, 

- while `==` is used as the **comparison operator** to compare two values.

```python
# Assign the variable
a = 5
a
```

```python
# Ask if it is equal to the value you think it should be
a == 5
```

There is one more boolean operator built into Python we have not mentioned yet. The keyword `in` allows you to check if strings are contained in larger strings. It can be used in other ways in Python, but for now we will use it to ask if the first string is a substring of the second string (if it appears in the other string).

Run the cells below for examples of `in`:

```python
"hello" in "hello world"
```

```python
"i" in "rhythm"
```

### And, Or and Not

`and`, `or`, and `not` are boolean operators that operate directly on booleans.

`and` returns True when both operands are True, False otherwise.

`or` returns False when both operands are False, True otherwise.

`not` gives back the opposite boolean.

These 3 operators are useful when writing code because they help control logic. You will find them useful when writing more complicated code in the course.

**Tip**: use parentheses to ensure your boolean operators are being applied correctly and without ambiguity.

Run the cells below to see these 3 operators in practice:

```python
(3 * 4 == 12) or (3 + 4 == 50)
```

```python
(3 * 4 == 12) and (3 + 4 == 50)
```

```python
not(3 + 4 == 50)
```

**Question 3.** Recall, the Pythagorean theorem is 

$$a^2 + b^2 = c^2$$

Using the Pythagorean theorem, set the variable `q3` equal to the length of the longest side of a right-angled triangle whose other side-lengths are `side_1` and `side_2`. Your code should not use the numbers 8 or 13 directly, but it should instead use the variables `side_1` and `side_2`.

**Hint:** `x**0.5` computes the square root of $x$.

```python
side_1 = 8
side_2 = 13

q3 = ...
q3
```

```python
grader.check("q3")
```

<!-- BEGIN QUESTION -->

**Question 4.** Assign the variable `q4` to a string that reads 

`"Lisa Jobs is 45 years old and lives in Canada."`

You **must** use the four variables defined below to create the string. This is similar to **Question 2**.

**Hint:** You will need to cast the variable age to an integer so that you can change it from 42 to 45.

```python
first_name = 'Lisa'
last_name = 'Jobs'
age = "42"
location = 'Canada'

q4 = ...
q4
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

When done exporting, download the .zip file by `SHIFT`-clicking on the file name and selecting **Save Link As**. Or, find the .zip file in the left side of the screen and right-click and select **Download**. You'll submit this .zip file for the assignment in Canvas to Gradescope for grading.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q3 (points: None)
```
>>> round(q3, 2) 
15.26
```
